---
title: "Road Traffic Accidents"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
    css: "css/font-size.css"
runtime: shiny_prerendered
description: >
  Get more practice with data visualisation and wrangling in the tidyverse: 
  practice visualisation with ggplot2, and data handling with the filter() and 
  mutate() functions.
---

<!-- Add JavaScript code for making the exercise code larger -->
<script language="JavaScript" src="js/exercise-font-size.js"></script>

```{r setup, include=FALSE}
# load packages ----------------------------------------------------------------

library(learnr)
library(gradethis)
library(tidyverse)
library(dsbox)

# set options for exercises and checking ---------------------------------------

gradethis_setup()

# hide non-exercise code chunks ------------------------------------------------

knitr::opts_chunk$set(echo = FALSE)
```

## Introduction

```{r biker, fig.margin = TRUE, echo = FALSE, out.width="100%", fig.cap="Photo by Carly Reeves on Unsplash"}
#might replace the image!
knitr::include_graphics("images/traffic.jpg")
```

In this tutorial, we will look at traffic accidents in Edinburgh.
The data are made available [online](https://data.gov.uk/dataset/cb7ae6f0-4be6-4935-9277-47e5ce24a11f/road-safety-data/datafile/36f1658e-b709-47e7-9f56-cca7aefeb8fe/preview) by the UK Government. 
It covers all recorded accidents in Edinburgh in 2018 and some of the variables were modified for the purposes of this tutorial. 

### Learning goals

- Continue practising data visualisation skills with ggplot2. 
- Filter data for certain attributes with `filter()`. 
- Create new variables based on existing variables in the data with `mutate()`. 

### Packages

We'll use the `tidyverse` package for the analysis and the `dsbox` package for the data.
These packages are already installed for you, so you load it as usual by running the following code:

```{r load-packages, exercise = TRUE}
library(tidyverse)
library(dsbox)
```

```{r load-packages-solution}
library(tidyverse)
library(dsbox)
```

```{r load-packages-check}
grade_this_code("The tidyverse and dsbox packages are now loaded!")
```

## Data

The data is in the `dsbox` package and it's called `accidents`. 

Below is an excerpt of the data dictionary. 
Note that it is long (there are lots of variables in the data), but we will be using a limited set of the variables for our analysis.

| Header                        |  Description
|:----------------|:--------------------------------
| `id` | Accident ID
| `easting` | Easting of accident location
| `northing` | Northing of accident location
| `longitude` | Longitude of accident location
| `latitude` | Latitude of accident location
| `police_force` | Police force
| `severity` | Accident severity: Fatal, Serious, Slight
| `vehicles` | Number of vehicles involved in accident
| `casualties` | Number of people injured in the accident
| `date` | Date of the accident
| `day_of_week` | Day of the week of the accident
| `time` | Time of the accident on the 24h clock
| `district` | Local authority district
| `highway` | Local authority highway
| `first_road_class` | Class of 1st road involved in accident: Motorway, A(M) road (A-road with motorway restrictions), A-road, B-road, C-road, Unclassified
| `first_road_number` | ID of 1st road (0 if unclassified)
| `road_type` | Type of road: Roundabout, One way street, Dual carriageway, Single carriageway, Slip road
| `speed_limit` | Speed limit on the road in mph
| `junction_detail` | Detail on junction where accident occurred: Crossroads, Mini-roundabout, More than 4 arms, Not within 20 metres of junction, Other junction, Private drive or entrance, Roundabout, Slip road, T or staggered junction
| `junction_control` | How junction was controlled: Authorised person, Auto traffic signal, Give way or uncontrolled, Missing / Out of range, Stop sign
| `second_road_class` | Class of 2st road involved in accident: A-road, B-road, C-road, Missing / Out of range, Motorway, Unclassified
| `second_road_number` | ID of 2nd road (0 if unclassified)
| `ped_cross_human` | Level of human control at a pedestrian crossing: Control by other authorised person, Control by school crossing patrol, None within 50 metres
| `ped_cross_physical` | Level of facilities controlling a pedestrian crossing: Central refuge, No physical crossing facilities within 50 metres, Non-junction crossing (pelican, puffin, toucan or similar light crossing), Pedestrian phase at traffic signal junction, Zebra crossing
| `light` | Light condition at the time of accident: Daylight, Darkness - lights lit, Darkness - lights unlit, Darkness - no lighting, Darkness - lighting unknown
| `weather` | Weather condition at the time of accident: Fine + no high winds, Raining + no high winds, Snowing + no high winds, Fine + high winds, Raining + high winds, Snowing + high winds, Fog or mist, Other, Unknown
| `road_surface` | Road surface conditions at the time of the accident: Dry, Wet or damp, Snow, Frost or ice, Flood over 3cm deep
| `special_condition` | Special condition at the site of the accident: None, Road sign or marking defective or obscured, Roadworks, Road surface defective
| `hazard` | Carriageway hazards: None, Other object on road, Previous accident, Pedestrian in carriageway - not injured
| `urban_rural` | Type of area the accident occurred in: 1 - urban, 2 - rural
| `police` | Did police officer attend the scene of the accident: No, No + accident self reported (using a self completion form), Yes

### First look at the data

You can take a peek at the data using the `glimpse()` function in the box below.

```{r glimpse-data, exercise = TRUE}
glimpse(accidents)
```

```{r quiz-rows}
question("What does each row in the dataset represent?",
    answer("The registration number of a car"),
    answer("The location of an accident"),
    answer("A recorded accident",
           correct = TRUE,
           message = "Each row in the dataset contains all information relating to an individual recorded accident in Edinburgh."),
    answer("An insurance claim "),
    allow_retry = TRUE
  )
```

How many accidents were recorded in Edinburgh in 2018? 
Use the following code chunk to submit your answer. 

```{r nrow, exercise = TRUE}

```

```{r nrow-hint-1}
Each row represents one recorded accident!
```

```{r nrow-hint-2}
Try using nrow()!
```

```{r nrow-check}
grade_this({
  if(identical(.result, 768) | identical(.result, 768L)) {
    pass("There are 768 rows, therefore, 768 accidents were recorded in Edinburgh in 2018.")
    }
  if(identical(.result, 31) | identical(.result, 31L)) {
    fail("Each observation is represented in one row. Did you calculate the number of columns instead of rows?")
    }
  fail("Not quite. Each observation is represented in one row. Try looking at the hints for some help!")
})
```

How many variables are recorded on these crashes? Use the code chunk below!

```{r ncol, exercise = TRUE}

```

```{r ncol-hint-1}
Each variable is displayed as a column.
```

```{r ncol-hint-2}
Try using ncol()!
```

```{r ncol-check}
grade_this({
  if(identical(.result, 31) | identical(.result, 31L)) {
    pass("Since there are 31 columns in the dataset, we know that 31 variables are recorded.")
  }
  if(identical(.result, 768) | identical(.result, 768L)) {
    fail("Each variable is recorded in a column. Did you maybe use the number of rows?")
  }
  fail("Not quite. Each variable is represented in a column. Try looking at the hints for some help!")
})
```

## Multi-vehicle accidents

How many accidents with 2 or more vehicles occurred in an urban area? 
Use the code chunk below to find out!

``` {r filter-accidents, exercise = TRUE}
___ %>%
  ___(___, ___) %>%
  nrow()
```
```{r filter-accidents-hint-1}
Use filter() to find the rows that match the criteria.
```

```{r filter-accidents-hint-2}
Review the data dictionary, specifically the variables urban_rural and vehicles.
```

``` {r filter-accidents-hint-3}
accidents %>%
  filter(vehicles >= ___, urban_rural == ___) %>%
  nrow()
```

```{r filter-accidents-check}
grade_this({
  if(identical(.result, 407) | identical(.result, 407L)) {
    pass("There are 407 rows that correspond to accidents with 2 or more vehicles occurred in an urban area.")
  } if(identical(.result, 72) | identical(.result, 72L)) {
    pass("Check which level corresponds to which type.")
  }
  fail("Take a peek at the hints!")
})
```

## Speed limits

Create a frequency table of the speed limits at which accidents happen (`speed_limit`).
Look at the hints for help!

```{r speed-frequency, exercise = TRUE}
___ %>%
  ___(___)
```

```{r speed-frequency-hint-1}
See the help for the `count()` function, specifically the 
`sort` argument for reporting the frequency table in descending order of counts, 
i.e. highest on top.
```

```{r speed-frequency-hint-2}
accidents %>%
  ___(___, sort = TRUE)
```


```{r speed-frequency-solution}
accidents %>%
  count(speed_limit, sort = TRUE)
```

```{r speed-frequency-check}
grade_this({
  if(identical(.result$n[1], 379L)) {
    pass("You have created the correct frequency table!")}
  fail("Not quite. See the hints for help!")
 })
```

```{r quiz-speed}
question("What is the most common speed limit in the
dataset?",
    answer("20", correct = TRUE),
    answer("30"),
    answer("40"),
    answer("50"),
    answer("60"),
    answer("70"),
    allow_retry = TRUE
  )
```

## Accident severity

### Visualising 

Recreate the following plot. 
To match the colours, you can use `scale_fill_viridis_d()`.

```{r example-plot, fig.width=7, fig.height=3, echo=FALSE}
ggplot(data = accidents, aes(x = severity, fill = light)) +
  geom_bar(position = "fill") +
  coord_flip() +
  labs(y = "Proportion", x = "Accident severity",
       fill = "Light condition", 
       title = "Light condition and accident severity") +
  scale_fill_viridis_d()
```

```{r plot-light, exercise = TRUE, fig.width=7, fig.height=3, echo = FALSE}
ggplot(data = ___, aes(x = ___, ___ = ___)) +
  geom____(___) +
  ___() +
  ___(y = ___, x = ___,
       ___ = ___, 
       title = ___)
```

```{r plot-light-hint-1}
ggplot(data = ___, aes(x = ___, ___ = ___)) +
  geom____(___) +
  ___() +
  ___(y = ___, x = ___,
       ___ = ___, 
       title = ___) +
  scale_fill_viridis_d()
```

```{r plot-light-hint-2}
ggplot(data = ___, aes(x = ___, fill = ___)) +
  geom_bar(___) +
  coord_flip() +
  labs(y = ___, x = ___,
       fill = ___, 
       title = ___) +
  scale_fill_viridis_d()
```

```{r plot-light-hint-3}
ggplot(data = ___, aes(x = ___, fill = ___)) +
  geom_bar(positions = ___) +
  coord_flip() +
  labs(y = ___, x = ___,
       fill = "Light condition", 
       title = ___) +
  scale_fill_viridis_d()
```

``` {r plot-light-solution}
ggplot(data = accidents, aes(x = severity, fill = light)) +
  geom_bar(position = "fill") +
  coord_flip() +
  labs(y = "Proportion", x = "Accident severity",
       fill = "Light condition", 
       title = "Light condition and accident severity") +
  scale_fill_viridis_d()
```

```{r plot-light-check}
grade_this_code("Well done!")
```

```{r plot-interpretation}
question("Which of the following are true? Check all that apply.", 
         answer("Most accidents occur in daylight",
                correct = TRUE),
         answer("Roughly 20 percent of serious accidents occurred in the darkness without lighting",
                message = "Look closely at the legend and the colours of the bars!"),
         answer("Crashes in the darkness tend to be more severe",
                correct = TRUE),
         answer("Fatal crashes have the highest proportion of crashes in the darkness where the lights are lit",
                message = "Compare the sizes of the segments for 'Darkness - lights lit' across the bars."),
         answer("Most slight accidents in the darkness happen without lighting."),
        allow_retry = TRUE
         )
```

### Customising labels

Recreate the same figure, but this time change the labels of the crash severity variable such that the dashes in the labels don't show up. 
There are many ways to do this, but in this tutorial, we'll focus on changing how the data are represented in the `light` variable using `mutate()`. 
Note that the colours in the figure might change, but that's ok.

```{r change-labels, exercise = TRUE}
accidents<- ___ %>%
  ___(___)

#now, copy the code from the previous exercise here!


```

``` {r change-labels-hint-1}
You could try using case_when()
```

```{r change-labels-hint-2}
accidents <- accidents %>%
  mutate(___ = case_when(___ == ___ ~ ___,
                            ...))
```

```{r change-labels-hint-3}
accidents <- accidents %>%
  mutate(light = case_when(
    light == "Daylight" ~ "Daylight",
    light == "Darkness - lights lit" ~ "Darkness, lights lit",
    ___
    ))
```

```{r change-labels-solution}
accidents <- accidents %>%
  mutate(light = case_when(
    light == "Daylight" ~ "Daylight",
    light == "Darkness - lights lit"       ~ "Darkness, lights lit",
    light == "Darkness - lights unlit"     ~ "Darkness, lights unlit", 
    light == "Darkness - no lighting"      ~ "Darkness, no lighting", 
    light == "Darkness - lighting unknown" ~ "Darkness, lighting unknown"
    ))

ggplot(data = accidents, aes(x = severity, fill = light)) +
  geom_bar(position = "fill") +
  coord_flip() +
  labs(y = "Proportion", x = "Accident severity",
       fill = "Light condition", 
       title = "Light condition and accident severity") +
  scale_fill_viridis_d()
```

```{r change-labels-check}
grade_this_code("Your solution is correct!")
```

## Wrap up

You have finished tutorial two, good job! 
We hope you enjoyed this lesson on data visualisation.
